# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #5</span></div>

### Due Monday, Oct. 31 at 2:00 pm $-$ upload your .zip submission to the PHYS 231 Canvas Gradebook.

***
## Lab Learning Outcomes:
By the end of the PHYS 231 Lab #5, students will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test a practical op-amp based inverting amplifier.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Experimentally explore the limitations of op-amp circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Use a log-log plot to examine the important characteristics of the inverting amplifier's frequency response.</span></b>

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Installing & Importing Packages</span></div>

## The 'InstallerCheck()' Function:

Run the 'InstallerCheck()' function contained within PHYS231.py by placing you're cursor inside the code cell below and then hitting 'Shift'+'Enter'.  If the function reports that some packages have been installed, run the cell with "InstallerCheck()" a second time.  It should report that all required packages are already installed.

```python
import PHYS231
PHYS231.InstallerCheck()
```

## Importing Packages/Modules:

Once the packages are installed, they can be imported using the notation:
```python
import packageName
```
where 'packageName' is the name of the package to be imported.  Execute the import statements below by placing your cursor within the code cell and hitting 'Shift' + 'Enter'.

```python
# Import required and commonly-used modules
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import uncertainties
from uncertainties import ufloat
    
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 5.ipynb")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Oscilloscope Measurement of the Amplifier Gain</span></div>

## NumPy Arrays

Data that you collect during the lab can be entered directly into a list using the following notation:

```python
x = [x0, x1, x2, x3, ..., xN]
y = [y0, y1, y2, y3, ..., yN]
```
These lists will contain $N+1$ elements. It's usually more convenient to enter data into a NumPy array rather than a list because you can do mathematical operations on arrays that you can't do on lists.  For example:

- squaring an array creates an new array with each of the elements from the original array squared
- multiplying two arrays of the same length creates a new array in which the each elements is the product of the corresponding elements from the original arrays
- ...

NumPy arrays are entered using the following syntax:

```python
x = np.array([x0, x1, x2, x3, ..., xN])
y = np.array([y0, y1, y2, y3, ..., yN])
```             

In the cell below, replace the "..." in NumPy array statements with your measurements of:

- frequency in kHz
- the amplitude of the signal from the function generator Vin in volts
- the uncertainty in Vin in volts
- the amplitude of the output voltage Vout in volts
- the uncertainty in Vout in volts

After executing the cell, the lengths of each of the arrays you've entered will be printed.  Make sure that all of the lengths are the same.

```python
# Enter your data into the arrays below.
# Please do NOT change the variable names.
frequency = np.array([...]) # kHz
Vin = np.array([...]) # V
errVin = np.array([...]) # V
Vout = np.array([...]) # V
errVout = np.array([...]) # V

# The statements below will print the lengths of the arrays.
print("Length of frequency array:", len(frequency))
print("Length of Vin array:", len(Vin))
print("Length of errVin array:", len(errVin))
print("Length of Vout array:", len(Vout))
print("Length of errVout array:", len(errVout))
```

## Pandas DataFrames

The Pandas package can be used to display your array data in a nicely-formated table.  The Pandas package creates "DataFrames" which are 2 dimensional data structures, like a 2 dimensional array, or a table with rows and columns.  The cell below creates and then displays a DataFrame using the arrays that you created in the cell above.

```python
# Create the data structure.
data = {
    "frequnecy (kHz)": frequency,
    "Vin (V)": Vin,
    "errVin (V)": errVin,
    "Vout (V)": Vout,
    "errVout (V)": errVout
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(data)
```

## $V_\mathrm{out}/V_\mathrm{in}$ versus $f$

In the cells below, calculate:
  
- $G = V_\mathrm{out}/V_\mathrm{in}$
- the uncertainty in $V_\mathrm{out}/V_\mathrm{in}$

Here, $G$ represents the absolute value of the gain of the amplifier.  We'll assume the relative uncertainty in the frequency is negligible compared the relative uncertainties of the other relevant parameters.

Recall propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

Please do **NOT** change the variable names.

```python
G = ...
```

```python
grader.check("G")
```

```python
# Note that, in Python, the square root can be called using np.sqrt(...).
errG = ...
```

```python
grader.check("G_uncertainty")
```

### Plot the data

The $V_\mathrm{out}/V_\mathrm{in}$ data can now be plotted as a function of frequency.  Execute the cell below to created a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

```python
fig1 = PHYS231.Scatter(
    xData = frequency,
    yData = G,
    yErrors = errG,
    xlabel = "frequency",
    ylabel = "G = Vout/Vin",
    xUnits = "kHz",
    yUnits = "",
)
```

## $V_\mathrm{out}/V_\mathrm{in}$ versus $f$

When measuring the response of a system (the gain of the amplifier, in this case) over a wide range of stimulus values (frequency, in this case) it is sometimes convenient to look at the data using *log-log plots*.  In this lab, such a plot will allow us to look for differences in the low-frequency and high-frequency response of the amplifier circuit.

In the cells below, we calculate the log of the the amplifier's gain ($G = V_\mathrm{out}/V_\mathrm{in}$) and the log of frequency.  All you need to do is execute the cells below.

```python
logf = np.log10(frequency)
logG = np.log10(G)
```

To calculate the error in the log of the gain, first note that if:
$$
y= \log x,
$$
Then:
$$
x = 10^y,
$$
such that:
$$
\ln x=\ln\left(10^y\right)=y\ln 10.
$$
Therefore, it must be the case that:
$$
y=\log x=\frac{1}{\ln 10}\ln x,
$$
which allows us to us propagation of errors to deduce that:
$$
\Delta y = \frac{1}{\ln 10}\left(\frac{\Delta x}{x}\right).
$$
In the cell below, enter an expression to calculate the uncertainty in the log of the gain.

```python
# Note that, in Python, np.log(x) represents ln(x) (natural logarithm, base e) and np.log10(x) represents log(x) (base 10 logarithm).
errlogG = ...
```

```python
grader.check("logG_uncertainty")
```

### Plot the $\log G$ vs $log f$

The $\log G$ data can now be plotted as a function of $\log f$.  Execute the cell below to created a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

```python
fig2 = PHYS231.Scatter(
    xData = logf,
    yData = logG,
    yErrors = errlogG,
    xlabel = "log f",
    ylabel = "log G",
    xUnits = "",
    yUnits = "",
)
```

## Linear Fit

You should have found that the high-frequency portion of your $\log G$ versus $\log f$ plot is linear with a negative slope.  The final bit of analysis that we'll do is to:

 - select the part of the data that is linear
 - find the slope of the linear section

First, we'll sort the data in the order of increasing values of $\log f$.

```python
inds = logf.argsort()
logf_sorted = logf[inds]
logG_sorted = logG[inds]
errlogG_sorted = errlogG[inds]

# Look at the contents of logf_sorted and logG_sorted.
logf_sorted, logG_sorted
```

Now select the part of the data that will used for the linear fit.  Adjust the start and stop values until the plot shows only the high-frequency linear portion of the data.

```python
# Set the limits.
start = ...
stop = ...

# Cut the data.
logf_cut = logf_sorted[start:stop]
logG_cut = logG_sorted[start:stop]
errlogG_cut = errlogG_sorted[start:stop]

# Plot the cut data.
fig3 = PHYS231.Scatter(
    xData = logf_cut,
    yData = logG_cut,
    yErrors = errlogG_cut,
    xlabel = "log f",
    ylabel = "log G",
    xUnits = "",
    yUnits = "",
)
```

## Linear Fits

The final thing that we'll do is fit a straight line through the cut data.  We'll use the 'LinearFit' contained in PHYS231.py to complete this task.  The function call to perform a weighted fit is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(t, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = t, yData = y, yErrors = erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
Execute the cell below to preform a weighted linear fit to your data.

For additional information about the 'LinearFit' function, see **Appendix B** near the end of this notebook. 

```python
m, b, dm, db, fig4 = PHYS231.LinearFit(
    xData = logf_cut,
    yData= logG_cut,
    yErrors = errlogG_cut,
    xlabel = "log f",
    ylabel = "log G",
    xUnits = "",
    yUnits = "",
)
```

<!-- BEGIN QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Feedback and Submission</span></div>

***
**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 231 Canvas gradebook.

Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/phys231/Python/images/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
PHYS231.LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS231.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS231.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

Last update: October 29, 2022

Tests for G (points: 1)
```
>>> len(G) != 0
True
```
```
>>> test = True
>>> for i in range(len(G)):
...     if np.isclose(G[i], Vout[i]/Vin[i], atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for G_uncertainty (points: 1)
```
>>> len(errG) != 0
True
```
```
>>> test = True
>>> for i in range(len(errG)):
...     if np.isclose(errG[i], np.sqrt((errVout[i]/Vin[i])**2 + (errVin[i]*Vout[i]/Vin[i]**2)**2), atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for logG_uncertainty (points: 1)
```
>>> len(errlogG) != 0
True
```
```
>>> test = True
>>> for i in range(len(errlogG)):
...     if np.isclose(errlogG[i], (1/np.log(10))*(errG[i]/G[i]), atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```
